# Loops

Topics:
1. `while` loops
2. `for` loops
<br>

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

For example, we could use `while` to count or to iterate over an array.

```julia

```

```julia

```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia

```

```julia

```

Now let's use `for` loops to create some addition tables, where the value of every entry is the sum of its row and column indices. <br>

Note that we iterate over this array via column-major loops in order to get the best performance. More information about fast indexing of multidimensional arrays inside nested loops can be found at https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1

First, we initialize an array with zeros.

```julia

```

```julia

```

Here's some syntactic sugar for the same nested `for` loop

```julia

```

```julia

```

The more "Julia" way to create this addition table would have been with an *array comprehension*.

```julia

```

### Exercises

#### 4.1 
Loop over integers between 1 and 100 and print their squares.

```julia

```

#### 4.2 
Add to the code above a bit to create a dictionary, `squares` that holds integers and their squares as key, value pairs such that

```julia
squares[10] == 100
```

```julia

```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia

```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```